# Enigma Machine Assignment
## Introduction
Please ensure you have read the instructions on Engage for a full introduction to this assignment.

This notebook contains the full instructions for each part of the assignment, but it assumes you already have an understanding of the assignment and the basics of how an Enigma machine works.

This notebook also includes instructions for how to format your code and some test cells like the ones you have seen in the exercise sheets. These are not exhaustive and outside of matching the names in the specified tests you are welcome to choose how you structure your code. We have tried to offer a good balance between being able to test your code, offering structure for the assignment, and also offering flexibility for you to be able to show off your own design ability.

The earlier stages of this assignment contain more guidance than the later ones. You are encouraged to research and use Python features that you have not been taught in the unit, and this can add to the quality of your code (as defined in the instructions on Engage). Remember also that advanced work which goes beyond the specification and is well documented will be rewarded (see the final section below).

## Part One – Simulation
In this part you will be producing a simulation of an Enigma machine using Python. Your code must be object-oriented, i.e. it must use classes. Some names of classes and methods will be specified in test cells, and you must match these where this is the case. Otherwise it is up to you how you implement the functionality and how you design your class structure.

**For all parts, you must write all of your Enigma machine code in separate files, i.e. not in the notebook cells.** In the cell below we have included an import statement `from enigma import *`. This will import the definitions (functions and classes) in a file called `enigma.py` which is located in the same directory, and a file is provided to get you started. If you wish to use a different file name, you can change this line and our tests will still run. Your code can be split across multiple files if you like, and you can include multiple import statements.

<!-- grade_id: cell-10268453ab40f7dd, points: 0, grade: false, solution: true -->
```python
from enigma import *
```

### Plug Leads
Let's start simulating the Enigma machine with the plugboard, specifically the leads. Each lead in the enigma machine connects two plugs in the plugboard. If we think of the functionality of the plugboard itself in terms of how it encodes a single character, the plugboard *aggregates* the leads, so it makes sense to make a class to represent the lead objects.

Write a class called `PlugLead`. The constructor should take a string of length two, which represents the two characters this lead should connect. So the following code:
```python3
lead = PlugLead("AG")
```
creates a lead which connects A to G. Remember leads are reversible so this lead also connects G to A. 

If you are still new to object-oriented design, when you are first conceptualising a class, you should take some time to think about what attributes it has and what methods it has. Some will be required by this specification, but you are always free to add more of your own.

As part of the specification, you must implement a method called `encode(c)` on your lead, which takes a single character `c`, and returns the result of this lead on this character. 

So, with the `lead` object we created above, `lead.encode("G")` should return `"A"`.

`lead.encode("D")` should return `"D"` – this lead had no impact on the letter D, it only connects A and G. 

Of course even though it would have no effect, it should not be possible to connect a letter to itself – there is only one plug for each letter on the plugboard. For this assignment you should write your code to be robust, obeying the physical limitations of the Engima machine (at least for part one; you may decide to lift these in your extension material). 

In general “robustness” is left up to your interpretation and discretion, but in this instance we will help out by suggesting your code should `raise` an error if someone tries to construct an invalid lead. You are welcome to include your own errors as custom classes, but in this case a `ValueError` would also be appropriate.

***Note:*** at this point we'll note that, of course, whenever the Enigma machine *encodes*, it also *decodes* for the exact same settings. We'll still use the method name `encode(…)` throughout our tests, and take it as assumed that the same method is used for encoding and decoding.

<!-- grade_id: cell-823cd8cc358d20ff, points: 5, grade: true, solution: false -->
```python
lead = PlugLead("AG")
assert(lead.encode("A") == "G")
assert(lead.encode("D") == "D")

lead = PlugLead("DA")
assert(lead.encode("A") == "D")
assert(lead.encode("D") == "A")
```

### Plugboard
Naturally we now need the plugboard itself to house our leads. 

An interesting part of object-oriented design is the idea that the *interface* of the object is really all an outsider needs to know. In the previous part we asked you to ensure a lead object supports `.encode("A")`, but how you achieve that is up to you.

Of course in this assignment we will be looking at your code, not just running automated tests. So two different implementations that both work could get different grades. But this is a programming unit and so it is part of the learning outcomes to learn how to write good code. In later units, you might be asked to show you are able to select and implement AI algorithms – not to test your ability to code, but to show you understand the techniques. Code quality is still important, it might even be worth marks, but it is not the primary goal.

With all that in mind, I am now going to ask you to write a class called `Plugboard`. The plugboard should accept leads which connect plugs via the method `.add(…)` as demonstrated below.

Like the leads, the plugboard should have an `.encode(…)` method as well, which should return the result of passing the character through the entire plugboard. 

These are the only interface requirements, but you are encouraged to elaborate on these with additional methods and/or constructor keywords. 

Remember that the Enigma machine only came with 10 leads to connect plugs. From now onwards, we will not always specifically point out the physical limitations that you need to model for robustness – you are expected to think of these and include them yourself. You do not need to know arcane details about how Enigma machines work, but from the information presented you can think of obvious incompatibilities and handle them.

<!-- grade_id: cell-aa9f93cfa0811bf2, points: 5, grade: true, solution: false -->
```python
plugboard = Plugboard()

plugboard.add(PlugLead("SZ"))
plugboard.add(PlugLead("GT"))
plugboard.add(PlugLead("DV"))
plugboard.add(PlugLead("KU"))

assert(plugboard.encode("K") == "U")
assert(plugboard.encode("A") == "A")
```

### Rotors
The number of possible combinations due to the plugboard is staggering, making brute force attempts to break a code extremely difficult. But using it alone would result in a simple substitution cipher – easily cracked with techniques like frequency analysis. The next step in the process, the rotors, allow the letter substitution to change mechanically every time a key is pressed, which prevents simple frequency analysis.

Over time, rotors with many different wiring patterns were developed, and different Enigma machines supported different types and numbers of rotors. We are not necessarily looking for exact historical accuracy in this assignment – for your extension material you might choose to be much more accurate or much less accurate! For this part, you should work on the following specification.

Your Enigma machine must support *three or four* rotors and one reflector – notice that a reflector is really just a type of rotor where the characters line up in 13 perfect pairs. The rotors will be numbered *from right to left*, which is the “path” the current takes when first entering the rotors. The signal goes through each rotor in turn, hits the reflector, then goes through the rotors again in reverse order (left to right). When it goes through the rotors in reverse order, it uses the *reverse* wiring. So if A is mapped to L when going from right to left, then L is mapped to A on the reverse journey (of course it is not possible to actually hit the same wire on the way back – a letter cannot encode into itself in the machine as a whole).

Ignoring reflectors, which never rotate, there are two types of rotating rotor, based on whether or not the rotor contains a *notch*. If the rotor contains a notch, then when it is on a certain position and is rotated, it will cause the rotor in the next slot to rotate as well. In addition, in a four slot Engima machine, the leftmost (fourth) rotor never rotates. Rotation will be explained in more detail shortly.

Each rotor can be chosen from a box containing seven possible wiring patterns. There are two rotors labelled `Beta` and `Gamma`. Then there are five rotors labelled with Roman numerals which do rotate: `I, II, III, IV, V`. You must also support three possible reflector wiring patterns, labelled `A, B, C`.

Note: the wiring patterns are all real, taken from [this page](https://en.wikipedia.org/wiki/Enigma_rotor_details). The page contains more rotors if you wish to consider them, or you can make up your own, but you must support the ones here. The following patterns all assume the rotors are in their default position, with their default ring setting, and going from right to left (the initial path).

<table><thead>
<tr><th></th><th colspan="26"><center>Mapping from letter</center></th></tr>
<tr><th style="text-align:left">Label</th><th>A</th><th>B</th><th>C</th><th>D</th><th>E</th><th>F</th><th>G</th><th>H</th><th>I</th><th>J</th><th>K</th><th>L</th><th>M</th><th>N</th><th>O</th><th>P</th><th>Q</th><th>R</th><th>S</th><th>T</th><th>U</th><th>V</th><th>W</th><th>X</th><th>Y</th><th>Z</th></tr></thead><tbody>
<tr><th style="text-align:left">Beta</th><td>L</td><td>E</td><td>Y</td><td>J</td><td>V</td><td>C</td><td>N</td><td>I</td><td>X</td><td>W</td><td>P</td><td>B</td><td>Q</td><td>M</td><td>D</td><td>R</td><td>T</td><td>A</td><td>K</td><td>Z</td><td>G</td><td>F</td><td>U</td><td>H</td><td>O</td><td>S</td></tr>
<tr><th style="text-align:left">Gamma</th><td>F</td><td>S</td><td>O</td><td>K</td><td>A</td><td>N</td><td>U</td><td>E</td><td>R</td><td>H</td><td>M</td><td>B</td><td>T</td><td>I</td><td>Y</td><td>C</td><td>W</td><td>L</td><td>Q</td><td>P</td><td>Z</td><td>X</td><td>V</td><td>G</td><td>J</td><td>D</td></tr>
<tr><th style="text-align:left">I</th><td>E</td><td>K</td><td>M</td><td>F</td><td>L</td><td>G</td><td>D</td><td>Q</td><td>V</td><td>Z</td><td>N</td><td>T</td><td>O</td><td>W</td><td>Y</td><td>H</td><td>X</td><td>U</td><td>S</td><td>P</td><td>A</td><td>I</td><td>B</td><td>R</td><td>C</td><td>J</td></tr>
<tr><th style="text-align:left">II</th><td>A</td><td>J</td><td>D</td><td>K</td><td>S</td><td>I</td><td>R</td><td>U</td><td>X</td><td>B</td><td>L</td><td>H</td><td>W</td><td>T</td><td>M</td><td>C</td><td>Q</td><td>G</td><td>Z</td><td>N</td><td>P</td><td>Y</td><td>F</td><td>V</td><td>O</td><td>E</td></tr>
<tr><th style="text-align:left">III</th><td>B</td><td>D</td><td>F</td><td>H</td><td>J</td><td>L</td><td>C</td><td>P</td><td>R</td><td>T</td><td>X</td><td>V</td><td>Z</td><td>N</td><td>Y</td><td>E</td><td>I</td><td>W</td><td>G</td><td>A</td><td>K</td><td>M</td><td>U</td><td>S</td><td>Q</td><td>O</td></tr>
<tr><th style="text-align:left">IV</th><td>E</td><td>S</td><td>O</td><td>V</td><td>P</td><td>Z</td><td>J</td><td>A</td><td>Y</td><td>Q</td><td>U</td><td>I</td><td>R</td><td>H</td><td>X</td><td>L</td><td>N</td><td>F</td><td>T</td><td>G</td><td>K</td><td>D</td><td>C</td><td>M</td><td>W</td><td>B</td></tr>
<tr><th style="text-align:left">V</th><td>V</td><td>Z</td><td>B</td><td>R</td><td>G</td><td>I</td><td>T</td><td>Y</td><td>U</td><td>P</td><td>S</td><td>D</td><td>N</td><td>H</td><td>L</td><td>X</td><td>A</td><td>W</td><td>M</td><td>J</td><td>Q</td><td>O</td><td>F</td><td>E</td><td>C</td><td>K</td></tr>
<tr><th style="text-align:left">A</th><td>E</td><td>J</td><td>M</td><td>Z</td><td>A</td><td>L</td><td>Y</td><td>X</td><td>V</td><td>B</td><td>W</td><td>F</td><td>C</td><td>R</td><td>Q</td><td>U</td><td>O</td><td>N</td><td>T</td><td>S</td><td>P</td><td>I</td><td>K</td><td>H</td><td>G</td><td>D</td></tr>
<tr><th style="text-align:left">B</th><td>Y</td><td>R</td><td>U</td><td>H</td><td>Q</td><td>S</td><td>L</td><td>D</td><td>P</td><td>X</td><td>N</td><td>G</td><td>O</td><td>K</td><td>M</td><td>I</td><td>E</td><td>B</td><td>F</td><td>Z</td><td>C</td><td>W</td><td>V</td><td>J</td><td>A</td><td>T</td></tr>
<tr><th style="text-align:left">C</th><td>F</td><td>V</td><td>P</td><td>J</td><td>I</td><td>A</td><td>O</td><td>Y</td><td>E</td><td>D</td><td>R</td><td>Z</td><td>X</td><td>W</td><td>G</td><td>C</td><td>T</td><td>K</td><td>U</td><td>Q</td><td>S</td><td>B</td><td>N</td><td>M</td><td>H</td><td>L</td></tr>
</tbody></table>

### Single Rotor Demonstration
From now on, the specification does not require you to follow specific class or method names, it is totally up to you.

Rotors still get bit more complicated when we introduce their settings and put multiple in the same machine. But you might want to take this opportunity to see if you can write some code to represent what we've seen so far.

In the cell below, you should demonstrate some basic rotor functionality using your classes. I have left some sample code in here in case you want to use it, either directly or for inspiration – I am not implying this is the only or best way to achieve it, and either way you will likely want to add to it. The idea is to show us that your rotors work with the concepts introduced so far.

This code cell is especially important if you end up struggling with the next part, to allow for partial credit.

<!-- grade_id: cell-ca471d294bbf30bd, points: 10, grade: true, solution: true -->
```python
class Rotor:
    """Superclass implementation of the rotors used in the Enigma Machine.

    Technical specifications of the Enigma rotors from:

    Sale, T.E., 2000. Technical specification of the Enigma [Online]. The Late Tony Sale's Codes and Ciphers Website
     (https://www.codesandciphers.org.uk/index.htm). Available from:
     https://www.codesandciphers.org.uk/enigma/rotorspec.htm [06 February 2022].
    """

    def __init__(self, ring_setting=1, position=1):
        """Rotor superclass constructor.

        :param ring_setting: The intended ring setting.
        :param position: The initial rotor position.
        """

        # will be overridden in subclasses
        # self.name = None
        self.rotor_number = int()
        self.ring_setting = ring_setting
        self.position = position
        self.encodings = list()
        self.encodings_rev = list()
        # would be logically easier to implement as int not str, but more abstract from Enigma construction.
        # A solution to this problem, is to enumerate self.turnover when required.
        # However, approach adds an order of time complexity, per rotor, to the Rotors.encode() algorithm.'''
        self.turnover = str

    def __str__(self):
        return self._name

    @staticmethod
    def get_reverse_encodings(encodings):
        """Helper function to generate the reverse encodings, as experienced by signal post-reflector.

        :param encodings: The standard encodings encodes from which to get the reverse encodings.
        :return: The reverse encodings.
        """

        import string
        import operator
        uc = list(string.ascii_uppercase)
        dict_to_swap = dict(zip(uc, encodings))
        rev_key_value = {value: key for (key, value) in dict_to_swap.items()}
        sorted_dict = dict(sorted(rev_key_value.items(), key=operator.itemgetter(0)))
        # get list of values
        dict_values = sorted_dict.values()
        reverse_encodings = list(dict_values)
        return reverse_encodings

    def set_rotor_number(self, rotor_num):
        """Sets the - from left to right - position in the rotor sub-system this rotor will occupy.

        I.e. From left to right, which position in the rotor system this rotor occupies.

        :param rotor_num: The position in the rotor sub-system.
        """

        self.rotor_number = rotor_num

    def rotate(self, positions=1):
        """Rotates the position of this rotor by param positions places.

        :param positions: The number of positions to rotate the position by - default is 1.
        """

        # edge case: invalid type, do not rotate
        if not isinstance(positions, int):
            pass
        # edge case: invalid value, do not rotate
        elif positions < 1:
            pass
        else:
            self.encodings = self.encodings[positions:] + self.encodings[0:positions]
            self.encodings_rev = self.encodings_rev[positions:] + self.encodings_rev[0:positions]
            self.position = (self.position + positions) % 26

    def get_name(self):
        """Gets the name of this rotor.

        :return: The name of this rotor.
        :rtype: str
        """

        return self._name

    def get_position(self):
        """Gets the current position value of this rotor.

        :return: The current position of this rotor.
        :rtype: str
        """

        return self.position

    def set_position(self, position):
        """Sets the rotational position for this rotor.

        :param position: The intended position.
        """

        # edge case: invalid type, do nothing
        if not isinstance(position, int):
            pass
        # edge case: invalid value, do nothing
        elif position < 1 or position > 26:
            pass
        # edge case: position == 1, no need to rotate
        elif position == 1:
            self.position = position
        else:
            self.position = position
            self.encodings = self.encodings[position - 1:] + self.encodings[0:position - 1]
            self.encodings_rev = self.encodings_rev[position - 1:] + self.encodings_rev[0:position - 1]

    def get_ring_setting(self):
        """Gets the current ring position.

        :return: The current ring position.
        """

        return self.ring_setting

    def set_ring_setting(self, ring_setting):
        """Sets the ring position.

        :param ring_setting: The intended ring setting for this encodings relative to 'A' MOD 26.
        """

        self.ring_setting = ring_setting
        # ring_setting of 1 means no adjustment...
        if ring_setting == 1:
            pass
        else:
            # ...therefore the actual adjustment needs to be reduced by 1
            self.rotate(ring_setting - 1)

    def encode(self, letter, reverse=False):
        letter = letter.upper()
        if ord(letter) < 65 or ord(letter) > 90:
            return ''
        else:
            if not reverse:
                relative_letter_value = ord(letter) - ord('A')

                return self.encodings[relative_letter_value]
            else:
                relative_letter_value = ord(letter) - ord('A')

                return self.encodings_rev[relative_letter_value]


class RotorI(Rotor):
    """Specialised Rotor; model I.

    """

    __encodings = ['E', 'K', 'M', 'F', 'L', 'G', 'D', 'Q', 'V', 'Z', 'N', 'T', 'O',
                   'W', 'Y', 'H', 'X', 'U', 'S', 'P', 'A', 'I', 'B', 'R', 'C', 'J']
    __encodings_rev = ['U', 'W', 'Y', 'G', 'A', 'D', 'F', 'P', 'V', 'Z', 'B', 'E', 'C',
                       'K', 'M', 'T', 'H', 'X', 'S', 'L', 'R', 'I', 'N', 'Q', 'O', 'J']
    __turnover = ['R']

    def __init__(self):
        super().__init__()
        self._name = 'I'
        super().__str__()
        self.encodings = RotorI.__encodings.copy()
        self.encodings_rev = RotorI.__encodings_rev.copy()
        self.turnover = RotorI.__turnover


class RotorII(Rotor):
    """Specialised Rotor; model II.

    """

    __encodings = ['A', 'J', 'D', 'K', 'S', 'I', 'R', 'U', 'X', 'B', 'L', 'H', 'W',
                   'T', 'M', 'C', 'Q', 'G', 'Z', 'N', 'P', 'Y', 'F', 'V', 'O', 'E']
    __encodings_rev = ['A', 'J', 'P', 'C', 'Z', 'W', 'R', 'L', 'F', 'B', 'D', 'K', 'O',
                       'T', 'Y', 'U', 'Q', 'G', 'E', 'N', 'H', 'X', 'M', 'I', 'V', 'S']
    __turnover = ['F']

    def __init__(self):
        super().__init__()
        self._name = 'II'
        super().__str__()
        self.encodings = RotorII.__encodings.copy()
        self.encodings_rev = RotorII.__encodings_rev.copy()
        self.turnover = RotorII.__turnover


class RotorIII(Rotor):
    """Specialised Rotor; model III.

    """

    __encodings = ['B', 'D', 'F', 'H', 'J', 'L', 'C', 'P', 'R', 'T', 'X', 'V', 'Z',
                   'N', 'Y', 'E', 'I', 'W', 'G', 'A', 'K', 'M', 'U', 'S', 'Q', 'O']
    __encodings_rev = ['T', 'A', 'G', 'B', 'P', 'C', 'S', 'D', 'Q', 'E', 'U', 'F', 'V',
                       'N', 'Z', 'H', 'Y', 'I', 'X', 'J', 'W', 'L', 'R', 'K', 'O', 'M']
    __turnover = ['W']

    def __init__(self):
        super().__init__()
        self._name = 'III'
        super().__str__()
        self.encodings = RotorIII.__encodings.copy()
        self.encodings_rev = RotorIII.__encodings_rev.copy()
        self.turnover = RotorIII.__turnover


class RotorIV(Rotor):
    """Specialised Rotor; model IV.

    """

    __encodings = ['E', 'S', 'O', 'V', 'P', 'Z', 'J', 'A', 'Y', 'Q', 'U', 'I', 'R',
                   'H', 'X', 'L', 'N', 'F', 'T', 'G', 'K', 'D', 'C', 'M', 'W', 'B']
    __encodings_rev = ['H', 'Z', 'W', 'V', 'A', 'R', 'T', 'N', 'L', 'G', 'U', 'P', 'X',
                       'Q', 'C', 'E', 'J', 'M', 'B', 'S', 'K', 'D', 'Y', 'O', 'I', 'F']
    __turnover = ['K']

    def __init__(self):
        super().__init__()
        self._name = 'IV'
        super().__str__()
        self.encodings = RotorIV.__encodings.copy()
        self.encodings_rev = RotorIV.__encodings_rev.copy()
        self.turnover = RotorIV.__turnover


class RotorV(Rotor):
    """Specialised Rotor; model V.

    """

    __encodings = ['V', 'Z', 'B', 'R', 'G', 'I', 'T', 'Y', 'U', 'P', 'S', 'D', 'N',
                   'H', 'L', 'X', 'A', 'W', 'M', 'J', 'Q', 'O', 'F', 'E', 'C', 'K']
    __encodings_rev = ['Q', 'C', 'Y', 'L', 'X', 'W', 'E', 'N', 'F', 'T', 'Z', 'O', 'S',
                       'M', 'V', 'J', 'U', 'D', 'K', 'G', 'I', 'A', 'R', 'P', 'H', 'B']
    __turnover = ['A']

    def __init__(self):
        super().__init__()
        self._name = 'I'
        super().__str__()
        self.encodings = RotorV.__encodings.copy()
        self.encodings_rev = RotorV.__encodings_rev.copy()
        self.turnover = RotorV.__turnover


class RotorVI(Rotor):
    """Specialised Rotor; model VI.

    """

    __encodings = ['J', 'P', 'G', 'V', 'O', 'U', 'M', 'F', 'Y', 'Q', 'B', 'E', 'N',
                   'H', 'Z', 'R', 'D', 'K', 'A', 'S', 'X', 'L', 'I', 'C', 'T', 'W']
    __encodings_rev = ['S', 'K', 'X', 'Q', 'L', 'H', 'C', 'N', 'W', 'A', 'R', 'V', 'G',
                       'M', 'E', 'B', 'J', 'P', 'T', 'Y', 'F', 'D', 'Z', 'U', 'I', 'O']
    __turnover = ['A', 'N']

    def __init__(self):
        super().__init__()
        self._name = 'VI'
        super().__str__()
        self.encodings = RotorVI.__encodings.copy()
        self.encodings_rev = RotorVI.__encodings_rev.copy()
        self.turnover = RotorVI.__turnover


class RotorVII(Rotor):
    """Specialised Rotor; model VII.

    """

    __encodings = ['N', 'Z', 'J', 'H', 'G', 'R', 'C', 'X', 'M', 'Y', 'S', 'W', 'B',
                   'O', 'U', 'F', 'A', 'I', 'V', 'L', 'P', 'E', 'K', 'Q', 'D', 'T']
    __encodings_rev = ['Q', 'M', 'G', 'Y', 'V', 'P', 'E', 'D', 'R', 'C', 'W', 'T', 'I',
                       'A', 'N', 'U', 'X', 'F', 'K', 'Z', 'O', 'S', 'L', 'H', 'J', 'B']
    __turnover = ['A', 'N']

    def __init__(self):
        super().__init__()
        self._name = 'VII'
        super().__str__()
        self.encodings = RotorVII.__encodings.copy()
        self.encodings_rev = RotorVII.__encodings_rev.copy()
        self.turnover = RotorVII.__turnover


class RotorVIII(Rotor):
    """Specialised Rotor; model VIII.

    """

    __encodings = ['F', 'K', 'Q', 'H', 'T', 'L', 'X', 'O', 'C', 'B', 'J', 'S', 'P',
                   'D', 'Z', 'R', 'A', 'M', 'E', 'W', 'N', 'I', 'U', 'Y', 'G', 'V']
    __encodings_rev = ['Q', 'J', 'I', 'N', 'S', 'A', 'Y', 'D', 'V', 'K', 'B', 'F', 'R',
                       'U', 'H', 'M', 'C', 'P', 'L', 'E', 'W', 'Z', 'T', 'G', 'X', 'O']
    __turnover = ['A', 'N']

    def __init__(self):
        super().__init__()
        self._name = 'VIII'
        super().__str__()
        self.encodings = RotorVIII.__encodings.copy()
        self.encodings_rev = RotorVIII.__encodings_rev.copy()
        self.turnover = RotorVIII.__turnover


class RotorBeta(Rotor):
    """Specialised Rotor; model Beta.

    """

    __encodings = ['L', 'E', 'Y', 'J', 'V', 'C', 'N', 'I', 'X', 'W', 'P', 'B', 'Q',
                   'M', 'D', 'R', 'T', 'A', 'K', 'Z', 'G', 'F', 'U', 'H', 'O', 'S']
    __encodings_rev = ['R', 'L', 'F', 'O', 'B', 'V', 'U', 'X', 'H', 'D', 'S', 'A', 'N',
                       'G', 'Y', 'K', 'M', 'P', 'Z', 'Q', 'W', 'E', 'J', 'I', 'C', 'T']
    __turnover = []

    def __init__(self):
        super().__init__()
        self._name = 'Beta'
        super().__str__()
        self.encodings = RotorBeta.__encodings.copy()
        self.encodings_rev = RotorBeta.__encodings_rev.copy()
        self.turnover = RotorBeta.__turnover


class RotorGamma(Rotor):
    """Specialised Rotor; model Gamma.

    """

    __encodings = ['F', 'S', 'O', 'K', 'A', 'N', 'U', 'E', 'R', 'H', 'M', 'B', 'T',
                   'I', 'Y', 'C', 'W', 'L', 'Q', 'P', 'Z', 'X', 'V', 'G', 'J', 'D']
    __encodings_rev = ['E', 'L', 'P', 'Z', 'H', 'A', 'X', 'J', 'N', 'Y', 'D', 'R', 'K',
                       'F', 'C', 'T', 'S', 'I', 'B', 'M', 'G', 'W', 'Q', 'V', 'O', 'U']
    __turnover = []

    def __init__(self):
        super().__init__()
        self._name = 'Gamma'
        super().__str__()
        self.encodings = RotorGamma.__encodings.copy()
        self.encodings_rev = RotorGamma.__encodings_rev.copy()
        self.turnover = RotorGamma.__turnover


class RotorAbstractFactory:
    """Rotor abstract factory - an implementation of the abstract factory design pattern.

    To enable dynamic creation of whatever flavor of Rotor, so that we can have Enigma machines with however many
    rotors.

    Abstract factory pattern implementation from:

    Chaudhary, M. 2021. Abstract Factory Method – Python Design Patterns [Online]. Uttar Pradesh: GeeksforGeeks.
     Available from: https://www.geeksforgeeks.org/abstract-factory-method-python-design-patterns/
     [Accessed Wed 23 Feb 2022].
    """

    def __init__(self, rotors_factory=None):
        """rotors_factory is the abstract factory

        :param rotors_factory: The Rotor type with which to configure the factory.
        """

        self.rotor_factory = rotors_factory

    def config_factory(self, rotors_factory: Rotor):
        """To enable the factory to be configured to create various specialised Rotor

        :param rotors_factory: The Rotor type with which to configure the factory.
        """

        self.rotor_factory = rotors_factory

    def create_rotor(self) -> Rotor:
        """Creates a specialised rotor (of type Rotor) polymorphically using the abstract factory.

        Flavour of Rotor is dependent on the current factory configuration.

        :return: The created Rotor.
        :rtype: Rotor
        """

        created_rotor = self.rotor_factory()
        return created_rotor


# create the rotor abstract factory
r_af = RotorAbstractFactory()
# configure the abstract factory for specialised Rotor of type RotorI
r_af.config_factory(RotorI)
# create the desired Rotor subclass
rotor = r_af.create_rotor()

assert(rotor.encode('A') == 'E')
assert (rotor.encode('A', True) == 'U')

```

### The Enigma Machine
To fully understand rotors, we need to imagine multiple of them in the Enigma machine itself. For this next part, you will need to model many more details of how the rotors work, and in addition work out how to incorporate them into a single machine that is capable of performing the full encoding path.

#### Multiple Rotors
For now, let's ignore the plugboard, and introduce the remaining details for the rotors.

It is common to see the selection of rotors specified in a single sequence from left to right, as the operator would see when looking down, such as on a [German code book](https://en.wikipedia.org/wiki/Enigma_machine#/media/File:Enigma_keylist_3_rotor.jpg). For example, on the top row of that linked image, you can see the rotors should be `I V III`. This means the first (rightmost) rotor is `III`, and so on. How you conceptualise the order inside your code is up to you providing it is consistent with the terminology here.

The code book also specifies each rotor's “ring setting” – in that image you can see the rotors on the top line should be set to `14 09 24` correspondingly. 

The rotor settings will be explained later, but let's first ensure you understand the way the wiring works with multiple rotors. We will show a worked example using three rotors, with all the settings in their default positions – this means you can read the character mappings directly from the table above.

Imagine a `III` rotor sat upright in the machine in the right hand position. The right hand side of the rotor has 26 *pins* and the left hand side has 26 *contacts*, one for each character. The pins on the right are connected to the contacts of the rotor housing which is wired into the plugboard. So if the plugboard sends a signal on the `A` contact of the housing, then this hits the `A` *pin* of the rotor, then this passes through the rotor's internal wiring (check the table for the `III` rotor) and we receive an output signal on the `B` *contact* of the rotor. This will now go into the next rotor: remember there are three or four rotors, followed by a reflector.

Suppose there is a `II` rotor in the middle position, a `I` rotor in the left position, and then a `B` reflector. Here is the full path when we send an `A` signal from the plugboard:
* `A` signal comes in
* `III` rotor receives signal on `A` pin, which connects to `B` contact
* `II` rotor receives signal on `B` pin, which connects to `J` contact
* `I` rotor receives signal on `J` pin, which connects to `Z` contact
* `B` reflector receives signal on `Z` connecting to `T` <br/>
  (now the signal goes backwards, hitting the *contacts* and coming out the *pins*)
* `I` rotor receives signal on `T` contact, which connects to `L` pin
* `II` rotor receives signal on `L` contact, which connects to `K` pin
* `III` rotor receives signal on `K` contact, which connects to `U` pin
* `U` signal is output

The final output for this `A` signal is a `U`. Make sure you can follow this using the table above, as things are about to get more complicated.

#### Rotation
As mentioned, the rightmost rotor (first in the order of the electrical circuit) rotates at *the start* of each keypress, i.e. *before* the character signal is passed through the circuit. This is what makes the Enigma machine more powerful than a fixed substitution cipher, one rotation causes a completely different circuit and substitution.

The rotation of the rotor advances a setting called the ***position*** of each rotor, which is visible through a window on the machine. This setting is marked on the rotor and is a character between `A` and `Z`, but it is helpful to think of this number as an *offset* rather than a *letter*. In the example above we assumed all of the rotors were set to position `A`. Rotating moves the *pins and the contacts* up by one position (`A` becomes `B`, etc).

We mentioned before that if we input `A` into the `III` rotor in its default position (which is labelled `A`) then it produces an output of `B`. In our full example we assumed that the rotors were set to `AAA`, but if we set them this way on the machine, then as soon as we press the `A` key the rightmost rotor would rotate giving `AAB`, and this is the circuit that would be made (rotation always happens first). Let's look at what happens in this setting.

Now if we input an `A` signal from the plugboard, it will come out of the `A` contact of the housing, but it will hit the `B` pin of the rotor due to its rotation by one position. The `B` *pin* is wired to the `D` *contact* (reading from the table).

***But*** the `D` contact has also been rotated one position inside the machine, so it actually lines up with the `C` *pin* of the next rotor in its default position. The other two rotors and the reflector work as normal, and on the way back rotor `II` sends a signal on its `E` pin. Since `III` is rotated, this hits the `F` contact which is wired to the `C` pin. But again since it is rotated, this his the `B` contact of the rotor housing, and is what is sent back to the plugboard.

Here is the full example again, now assuming the rotors are set to `AAB` instead of `AAA`:
* `A` signal comes in
* `III (B)` rotor receives signal on `B` pin, which connects to `D` contact 
* `II (A)` rotor receives signal on `C` pin, which connects to `D` contact
* `I (A)` rotor receives signal on `D` pin, which connects to `F` contact
* `B` reflector receives signal on `F` connecting to `S`
* `I (A)` rotor receives signal on `S` contact, which connects to `S` pin
* `II (A)` rotor receives signal on `S` contact, which connects to `E` pin
* `III (B)` rotor receives signal on `F` contact, which connects to `C` pin
* `B` signal is output

You can try using the Enigma machine emulator [on this page](https://www.101computing.net/enigma-machine-emulator/). It defaults to the same settings: rotors `I II III` all initially set to position `A`, so when you press the `A` key on the keyboard you should get `B`.

Make sure you can follow this path using the table of wirings to help you understand how you will implement the behaviour. If you simply do not follow, there are lots of explanations and videos online for Enigma machines. 

#### Ring Settings
In addition each rotor could be configured by changing the *ring setting*, which is a fixed offset that would apply between the internal wiring and the external markings. The ring settings were either given from `A-Z` or `01-26` – you saw the latter in the code book image linked above, and we'll use these too to avoid confusion with the *position* setting.

If a rotor's ring setting is set to `01` then nothing is changed, the wiring is exactly as written in the table above.

*Increasing* the ring setting has the exact same effect as *decreasing* the position setting. It shifts the internal wiring in the opposite direction.

Earlier we said `A` becomes `U` with the given rotors actually set to `AAA` – we'd have to start on `AAZ` to get this result on the machine, since the rotation happens first (try it on the emulator). Alternatively, we could set the ring position to `02` on the rightmost rotor and set the initial positions to `AAA`, and we'll get the same result for a single press (again you can try this on the emulator; it uses letters for ring settings, if you click the rotor you can set the ring setting to `B`).

If we keep pressing `A`, the two configurations will produce many of the same characters, but not always. We have detailed how the rightmost rotor rotates, but not the others, and this detail will eventually produce a difference between the two sets of settings above.

#### Turnover
The rotors labelled `I` to `V` have *notches*. If a rotor has a notch and is currently set to its notch position, then it will turn the next rotor on the next keypress (this is called *turnover*). Here are the notch positions:

<table><thead><tr><th>Rotor</th><th>Notch</th></tr></thead><tbody><tr><td style="text-align:center">I</td><td style="text-align:center">Q</td></tr><tr><td style="text-align:center">II</td><td style="text-align:center">E</td></tr><tr><td style="text-align:center">III</td><td style="text-align:center">V</td></tr><tr><td style="text-align:center">IV</td><td style="text-align:center">J</td></tr><tr><td style="text-align:center">V</td><td style="text-align:center">Z</td></tr></tbody></table>

So if a `II` rotor in the first, rightmost, slot of the machine is *currently set to* position `E` then when you press a key the first rotor will turn to position `F` *and* the rotor in the second position will turn as well, and then the electrical signal will be sent through the circuit.

If the `II` rotor had been in the second position, then it will obviously turn much more slowly. But if its position is set to `E` and a key is pressed then it will rotate and turn the rotor in the third position also. 

There is an important detail here called the *double step*. Normally the second rotor will only turn once every 26 turns of the rightmost rotor. But if the second rotor is *on* its notch setting, then it will turn again *as it turns the third rotor*. Obviously the second rotor must have just rotated to land on its notch, so it actually rotates for two keypresses in a row.

Suppose we have the rotors `I II III`, on positions `A C U`. Pressing a key turns the `III` rotor and we get `A C V`. Now the `III` rotor is on its notch, so pressing a key also turns `II` and we get `A D W`. If we continue pressing keys we'll get `A D X`, `A D Y`, `A D Z`, `A D A`, and so on. Several keypresses later we wrap round again and approach the notch on `III` again on setting `A D U`. When we press once we get `A D V`. Now `III` is on its notch so pressing again turns `II` and we get `A E W`. But now `II` is on its notch, so when we press again *all three* rotors rotate and we get `B F X` – `III` turns because it always turns, `II` turns because it is on its notch, and `I` turns because `II` turned on its notch (turnover).

Notice the *ring setting* is inconsequential in this turnover process – it only requires the current position to line up with the notch setting.

The four-rotor machines did not have a additional lever, and so whether the third rotor had a notch or not the fourth rotor would not turn. In addition if a notchless rotor (e.g. `Beta`) was in the first position, then it will never cause the second rotor to rotate. The rightmost rotor will still always rotate exactly once on every keypress. Notchless rotors can still be set to different position settings as part of the setup process.

If you are curious about the mechanism, you can see a video of a mock-up version of an Enigma machine in action [here](https://www.youtube.com/watch?v=hcVhQeZ5gI4), showing how the the ratchets, levers, and notches contribute towards the rotation on each keypress, and also demonstrating the double step (around 26 seconds into the video).

### Multiple Rotor Demonstration
You have free reign to model the rotors however you wish, and you are encouraged to think about how object-oriented design principles and features might apply.

In the cell below, demonstrate that your rotors work. The code is left entirely up to you, though you must ensure to demonstrate the following:
* With rotors `I II III`, reflector `B`, ring settings `01 01 01`, and initial positions `A A Z`, encoding an `A` produces a `U`.
* With rotors `I II III`, reflector `B`, ring settings `01 01 01`, and initial positions `A A A`, encoding an `A` produces a `B`.
* With rotors `I II III`, reflector `B`, ring settings `01 01 01`, and initial positions `Q E V`, encoding an `A` produces an `L`.
* With rotors `IV V Beta`, reflector `B`, ring settings `14 09 24`, and initial positions `A A A`, encoding an `H` produces a `Y`.
* With rotors `I II III IV`, reflector `C`, ring settings `07 11 15 19`, and initial positions `Q E V Z`, encoding a `Z` produces a `V`.

<!-- grade_id: cell-2770e919610e99e7, points: 10, grade: true, solution: true -->
```python
from abc import ABC
from typing import Any
import string


def rotate_letter(letter, rotation):
    """Rotates [uppercase] characters around the alphabet.  Works in both directions.

    :param letter: The letter to be rotations.
    :param rotation: The number of positions to rotate.
    :return: The rotations letter.
    """

    # invalid input
    if len(letter) != 1:
        return letter
    # letter is not in alphabet
    elif not letter.isalpha():
        return letter
    else:
        letter = letter.upper()
        ordv = ord(letter)
        ans = ((ordv - 65 + rotation) % 26) + 65
        return chr(ans)


class PlugLead:
    def __init__(self, patch):
        patch = patch.upper()
        self.mappings = dict()
        self.mappings = {patch[0]: patch[1], patch[1]: patch[0]}

    def encode(self, letter):
        if letter.upper() in self.mappings:
            return self.mappings[letter]
        else:
            return letter.upper()


class Plugboard:
    def __init__(self):
        uppercase_letters_string = string.ascii_uppercase
        uppercase_letters_list = list(uppercase_letters_string)
        self.mappings = {k: v for k, v in zip(uppercase_letters_list, uppercase_letters_list)}
        self.patchings = dict()
        self.plugleads = list()
        self.__max_num_of_plugleads = 10
        self.__num_of_plugleads = 0

    def add(self, pluglead: PlugLead):
        """Add one of ten available PlugLeads to the Plugboard.

        Patches two letters specified in the name parameter of the PlugLead.  E.g. the PlugLead argument 'AB' would
        patch the letters 'A' abd 'B', so that 'A' would be encoded to 'B' and 'B' would be encoded to 'A'.

        :param pluglead: A PlugLead object specifying the two letters for a bidirectional patch.
        :type pluglead: PlugLead
        """

        # edge case: attempt to make a patch that utilises an already patched letter
        for i in pluglead.mappings:
            for j in self.patchings:
                if i == j:
                    raise ValueError
        # edge case: attempt to add more than 10 PlugLeads
        if len(self.plugleads) == self.__max_num_of_plugleads:
            raise ValueError
        else:
            self.plugleads.append(pluglead)
            self.mappings.update(pluglead.mappings)
            self.patchings.update(pluglead.mappings)
            self.__num_of_plugleads += 1

    # @TODO imp. un_patch: delete from self.plugleads; reset self.mappings
    def un_patch(self, pluglead: PlugLead):
        if not isinstance(pluglead, PlugLead):
            pass
        elif pluglead not in self.plugleads:
            pass
        else:
            self.plugleads.remove(pluglead)

    def encode(self, letter):
        """Encodes param letter

        Encoding of the letter depends on the current state of the plugboard. I.e. what - if any - patches have beem
        made.

        :param letter: the letter to be encoded
        :return: the encoded letter
        """
        return self.mappings[letter]


class Rotors:
    """

    """

    def __init__(self, ):
        self.rotors = list()
        self.num_of_rotors = 0

    def add_rotor_to_rotors(self, new_rotor):
        self.rotors.append(new_rotor)
        self.num_of_rotors += 1
        self.rotors[self.num_of_rotors - 1].set_rotor_number(self.num_of_rotors)

    def get_num_of_rotors(self):
        return self.num_of_rotors

    def get_rotor_name(self, rotor_number):
        # @TODO get name of encodings from intended list of Rotors, via param rotor_number
        pass

    def encode(self, letter, reverse=False):
        """Encodes an inputted letter.

        - Uses the specific Enigma encodings map to encode a letter.
        - Checks if turnover position has been reached, if so triggers rotation of next encodings.

        :param letter: The letter to encode.
        :param reverse: True for standard encoding (default), False for reverse encoding.
        :return: The encoded letter
        """

        letter = letter.upper()

        if not reverse:
            # rotate the right-most rotor
            self.rotors[2].rotate()

            # check for rotor 2 turnover
            r2_position_letter = Enigma.enalpharate_position(self.rotors[2].position)
            if r2_position_letter in self.rotors[2].turnover:
                self.rotors[1].rotate()

            # check for rotor 1 turnover
            r1_position_letter = Enigma.enalpharate_position(self.rotors[1].position)
            if r1_position_letter in self.rotors[1].turnover:
                self.rotors[0].rotate()

            # encode letter by rotor 3
            letter = self.rotors[2].encode(letter)
            print(f'Rotor {self.rotors[-1].rotor_number} ({self.rotors[2].__str__()}) encoded: {letter}')

            # adjust next input by rotation, relative to initial setting
            if self.rotors[2].position != 1:
                rotation = 27 - self.rotors[2].position
                letter = rotate_letter(letter, rotation)
                print(f'Rotor {self.rotors[2].rotor_number} rotated, so input to next is: {letter}')

            # encode letter by rotor 2
            letter = self.rotors[1].encode(letter)
            print(f'Rotor {self.rotors[1].rotor_number} ({self.rotors[1].__str__()}) encoded: {letter}')

            # adjust next input by rotation, relative to initial setting
            if self.rotors[1].position != 1:
                rotation = 27 - self.rotors[1].position
                letter = rotate_letter(letter, rotation)
                print(f'Rotor {self.rotors[1].rotor_number} rotated, so input to next is: {letter}')

            # encode letter by rotor 1
            letter = self.rotors[0].encode(letter)
            print(f'Rotor {self.rotors[0].rotor_number} ({self.rotors[0].__str__()}) encoded: {letter}')

            # adjust next input by rotation, relative to initial setting
            if self.rotors[0].position != 1:
                rotation = 27 - self.rotors[0].position
                letter = rotate_letter(letter, rotation)
                print(f'Rotor {self.rotors[0].rotor_number} rotated, so input to next is: {letter}')
        else:
            # encode letter by rotor 1
            letter = self.rotors[0].encode(letter, True)
            print(f'Rotor {self.rotors[0].rotor_number} ({self.rotors[0].__str__()}) encoded: {letter}')

            if self.rotors[0].position != 1:
                rotation = 27 - self.rotors[0].position
                letter = rotate_letter(letter, rotation)
                print(f'Rotor {self.rotors[0].rotor_number} rotated, so input to next is: {letter}')

            # encode letter by rotor 2
            letter = self.rotors[1].encode(letter, True)
            print(f'Rotor {self.rotors[1].rotor_number} ({self.rotors[1].__str__()}) encoded: {letter}')

            # adjust next input by rotation, relative to initial setting
            if self.rotors[1].position != 1:
                rotation = 27 - self.rotors[1].position
                letter = rotate_letter(letter, rotation)
                print(f'Rotor {self.rotors[1].rotor_number} rotated, so input to next is: {letter}')

            # encode letter by rotor 3
            letter = self.rotors[2].encode(letter, True)
            print(f'Rotor {self.rotors[-1].rotor_number} ({self.rotors[2].__str__()}) encoded: {letter}')

            # adjust next input by rotation, relative to initial setting
            if self.rotors[2].position != 1:
                rotation = 27 - self.rotors[2].position
                letter = rotate_letter(letter, rotation)
                print(f'Rotor {self.rotors[2].rotor_number} rotated, so output is: {letter}')

        return letter

    @staticmethod
    def rotate_letter(letter, rotation):
        """Rotates [uppercase] characters around the alphabet.  Works in both directions.

        :param letter: The letter to be rotations.
        :param rotation: The number of positions to rotate.
        :return: The rotations letter.
        """

        # invalid input
        if len(letter) != 1:
            return letter
        # letter is not in alphabet
        elif not letter.isalpha():
            return letter
        else:
            letter = letter.upper()
            ordv = ord(letter)
            ans = ((ordv - 65 + rotation) % 26) + 65
            return chr(ans)


class Rotor:
    """Superclass implementation of the rotors used in the Enigma Machine.

    Technical specifications of the Enigma rotors from:

    Sale, T.E., 2000. Technical specification of the Enigma [Online]. The Late Tony Sale's Codes and Ciphers Website
     (https://www.codesandciphers.org.uk/index.htm). Available from:
     https://www.codesandciphers.org.uk/enigma/rotorspec.htm [06 February 2022].
    """

    def __init__(self, ring_setting=1, position=1):
        """Rotor superclass constructor.

        :param ring_setting: The intended ring setting.
        :param position: The initial rotor position.
        """

        # will be overridden in subclasses
        # self.name = None
        # @TODO this is not working, the individual rotors do not have correct name
        self.rotor_number = int()
        self.ring_setting = ring_setting
        self.position = position
        self.encodings = list()
        self.encodings_rev = list()
        # would be logically easier to implement as int not str, but more abstract from Enigma construction.
        # A solution to this problem, is to enumerate self.turnover when required.
        # However, approach adds an order of time complexity, per rotor, to the Rotors.encode() algorithm.'''

        self.input_offset = 0
        self.output_offset = 0

    def __str__(self):
        return self._name

    @staticmethod
    def get_reverse_encodings(encodings):
        """Helper function to generate the reverse encodings, as experienced by signal post-reflector.

        :param encodings: The standard encodings encodes from which to get the reverse encodings.
        :return: The reverse encodings.
        """

        import string
        import operator
        uc = list(string.ascii_uppercase)
        dict_to_swap = dict(zip(uc, encodings))
        rev_key_value = {value: key for (key, value) in dict_to_swap.items()}
        sorted_dict = dict(sorted(rev_key_value.items(), key=operator.itemgetter(0)))
        # get list of values
        dict_values = sorted_dict.values()
        reverse_encodings = list(dict_values)
        return reverse_encodings

    def set_rotor_number(self, rotor_num):
        """Sets the - from left to right - position in the rotor sub-system this rotor will occupy.

        I.e. From left to right, which position in the rotor system this rotor occupies.

        :param rotor_num: The position in the rotor sub-system.
        """

        self.rotor_number = rotor_num

    def rotate(self, positions=1):
        """Rotates the position of this rotor by param positions places.

        :param positions: The number of positions to rotate the position by - default is 1.
        """

        # edge case: invalid type, do not rotate
        if not isinstance(positions, int):
            pass
        # edge case: invalid value, do not rotate
        elif positions < 1:
            pass
        else:
            self.encodings = self.encodings[positions:] + self.encodings[0:positions]
            self.encodings_rev = self.encodings_rev[positions:] + self.encodings_rev[0:positions]
            self.position = (self.position + positions) % 26

    def get_name(self):
        """Gets the name of this rotor.

        :return: The name of this rotor.
        :rtype: str
        """

        return self._name

    def get_position(self):
        """Gets the current position value of this rotor.

        :return: The current position of this rotor.
        :rtype: str
        """

        return self.position

    def set_position(self, position):
        """Sets the rotational position for this rotor.

        :param position: The intended position.
        """

        # edge case: invalid type, do nothing
        if not isinstance(position, int):
            pass
        # edge case: invalid value, do nothing
        elif position < 1 or position > 26:
            pass
        # edge case: position == 1, no need to rotate
        elif position == 1:
            self.position = position
        else:
            self.position = position
            self.encodings = self.encodings[position - 1:] + self.encodings[0:position - 1]
            self.encodings_rev = self.encodings_rev[position - 1:] + self.encodings_rev[0:position - 1]

    def get_ring_setting(self):
        """Gets the current ring position.

        :return: The current ring position.
        """

        return self.ring_setting

    def set_ring_setting(self, ring_setting):
        """Sets the ring position.

        :param ring_setting: The intended ring setting for this encodings relative to 'A' MOD 26.
        """

        self.ring_setting = ring_setting
        # ring_setting of 1 means no adjustment...
        if ring_setting == 1:
            pass
        else:
            # ...therefore the actual adjustment needs to be reduced by 1
            self.rotate(-ring_setting)
            print()

    def encode(self, letter, reverse=False):
        """Encodes a letter based upon this rotor's mappings.

        :param letter: The letter to be encoded.
        :param reverse: Perform a reverse encoding post-reflector.
        :return: The encoded letter.
        """

        letter = letter.upper()
        if ord(letter) < 65 or ord(letter) > 90:
            return ''
        else:
            if not reverse:
                relative_letter_value = ord(letter) - ord('A')
                return self.encodings[relative_letter_value]
            else:
                relative_letter_value = ord(letter) - ord('A')
                return self.encodings_rev[relative_letter_value]


class RotorI(Rotor):
    """Specialised Rotor; model I.

    """

    __encodings = ['E', 'K', 'M', 'F', 'L', 'G', 'D', 'Q', 'V', 'Z', 'N', 'T', 'O',
                   'W', 'Y', 'H', 'X', 'U', 'S', 'P', 'A', 'I', 'B', 'R', 'C', 'J']
    __encodings_rev = ['U', 'W', 'Y', 'G', 'A', 'D', 'F', 'P', 'V', 'Z', 'B', 'E', 'C',
                       'K', 'M', 'T', 'H', 'X', 'S', 'L', 'R', 'I', 'N', 'Q', 'O', 'J']
    __turnover = ['R']

    def __init__(self):
        super().__init__()
        self._name = 'I'
        super().__str__()
        self.encodings = RotorI.__encodings.copy()
        self.encodings_rev = RotorI.__encodings_rev.copy()
        self.turnover = RotorI.__turnover


class RotorII(Rotor):
    """Specialised Rotor; model II.

    """

    __encodings = ['A', 'J', 'D', 'K', 'S', 'I', 'R', 'U', 'X', 'B', 'L', 'H', 'W',
                   'T', 'M', 'C', 'Q', 'G', 'Z', 'N', 'P', 'Y', 'F', 'V', 'O', 'E']
    __encodings_rev = ['A', 'J', 'P', 'C', 'Z', 'W', 'R', 'L', 'F', 'B', 'D', 'K', 'O',
                       'T', 'Y', 'U', 'Q', 'G', 'E', 'N', 'H', 'X', 'M', 'I', 'V', 'S']
    __turnover = ['F']

    def __init__(self):
        super().__init__()
        self._name = 'II'
        super().__str__()
        self.encodings = RotorII.__encodings.copy()
        self.encodings_rev = RotorII.__encodings_rev.copy()
        self.turnover = RotorII.__turnover


class RotorIII(Rotor):
    """Specialised Rotor; model III.

    """

    __encodings = ['B', 'D', 'F', 'H', 'J', 'L', 'C', 'P', 'R', 'T', 'X', 'V', 'Z',
                   'N', 'Y', 'E', 'I', 'W', 'G', 'A', 'K', 'M', 'U', 'S', 'Q', 'O']
    __encodings_rev = ['T', 'A', 'G', 'B', 'P', 'C', 'S', 'D', 'Q', 'E', 'U', 'F', 'V',
                       'N', 'Z', 'H', 'Y', 'I', 'X', 'J', 'W', 'L', 'R', 'K', 'O', 'M']
    __turnover = ['W']

    def __init__(self):
        super().__init__()
        self._name = 'III'
        super().__str__()
        self.encodings = RotorIII.__encodings.copy()
        self.encodings_rev = RotorIII.__encodings_rev.copy()
        self.turnover = RotorIII.__turnover


class RotorIV(Rotor):
    """Specialised Rotor; model IV.

    """

    __encodings = ['E', 'S', 'O', 'V', 'P', 'Z', 'J', 'A', 'Y', 'Q', 'U', 'I', 'R',
                   'H', 'X', 'L', 'N', 'F', 'T', 'G', 'K', 'D', 'C', 'M', 'W', 'B']
    __encodings_rev = ['H', 'Z', 'W', 'V', 'A', 'R', 'T', 'N', 'L', 'G', 'U', 'P', 'X',
                       'Q', 'C', 'E', 'J', 'M', 'B', 'S', 'K', 'D', 'Y', 'O', 'I', 'F']
    __turnover = ['K']

    def __init__(self):
        super().__init__()
        self._name = 'IV'
        super().__str__()
        self.encodings = RotorIV.__encodings.copy()
        self.encodings_rev = RotorIV.__encodings_rev.copy()
        self.turnover = RotorIV.__turnover


class RotorV(Rotor):
    """Specialised Rotor; model V.

    """

    __encodings = ['V', 'Z', 'B', 'R', 'G', 'I', 'T', 'Y', 'U', 'P', 'S', 'D', 'N',
                   'H', 'L', 'X', 'A', 'W', 'M', 'J', 'Q', 'O', 'F', 'E', 'C', 'K']
    __encodings_rev = ['Q', 'C', 'Y', 'L', 'X', 'W', 'E', 'N', 'F', 'T', 'Z', 'O', 'S',
                       'M', 'V', 'J', 'U', 'D', 'K', 'G', 'I', 'A', 'R', 'P', 'H', 'B']
    __turnover = ['A']

    def __init__(self):
        super().__init__()
        self._name = 'I'
        super().__str__()
        self.encodings = RotorV.__encodings.copy()
        self.encodings_rev = RotorV.__encodings_rev.copy()
        self.turnover = RotorV.__turnover


class RotorVI(Rotor):
    """Specialised Rotor; model VI.

    """

    __encodings = ['J', 'P', 'G', 'V', 'O', 'U', 'M', 'F', 'Y', 'Q', 'B', 'E', 'N',
                   'H', 'Z', 'R', 'D', 'K', 'A', 'S', 'X', 'L', 'I', 'C', 'T', 'W']
    __encodings_rev = ['S', 'K', 'X', 'Q', 'L', 'H', 'C', 'N', 'W', 'A', 'R', 'V', 'G',
                       'M', 'E', 'B', 'J', 'P', 'T', 'Y', 'F', 'D', 'Z', 'U', 'I', 'O']
    __turnover = ['A', 'N']

    def __init__(self):
        super().__init__()
        self._name = 'VI'
        super().__str__()
        self.encodings = RotorVI.__encodings.copy()
        self.encodings_rev = RotorVI.__encodings_rev.copy()
        self.turnover = RotorVI.__turnover


class RotorVII(Rotor):
    """Specialised Rotor; model VII.

    """

    __encodings = ['N', 'Z', 'J', 'H', 'G', 'R', 'C', 'X', 'M', 'Y', 'S', 'W', 'B',
                   'O', 'U', 'F', 'A', 'I', 'V', 'L', 'P', 'E', 'K', 'Q', 'D', 'T']
    __encodings_rev = ['Q', 'M', 'G', 'Y', 'V', 'P', 'E', 'D', 'R', 'C', 'W', 'T', 'I',
                       'A', 'N', 'U', 'X', 'F', 'K', 'Z', 'O', 'S', 'L', 'H', 'J', 'B']
    __turnover = ['A', 'N']

    def __init__(self):
        super().__init__()
        self._name = 'VII'
        super().__str__()
        self.encodings = RotorVII.__encodings.copy()
        self.encodings_rev = RotorVII.__encodings_rev.copy()
        self.turnover = RotorVII.__turnover


class RotorVIII(Rotor):
    """Specialised Rotor; model VIII.

    """

    __encodings = ['F', 'K', 'Q', 'H', 'T', 'L', 'X', 'O', 'C', 'B', 'J', 'S', 'P',
                   'D', 'Z', 'R', 'A', 'M', 'E', 'W', 'N', 'I', 'U', 'Y', 'G', 'V']
    __encodings_rev = ['Q', 'J', 'I', 'N', 'S', 'A', 'Y', 'D', 'V', 'K', 'B', 'F', 'R',
                       'U', 'H', 'M', 'C', 'P', 'L', 'E', 'W', 'Z', 'T', 'G', 'X', 'O']
    __turnover = ['A', 'N']

    def __init__(self):
        super().__init__()
        self._name = 'VIII'
        super().__str__()
        self.encodings = RotorVIII.__encodings.copy()
        self.encodings_rev = RotorVIII.__encodings_rev.copy()
        self.turnover = RotorVIII.__turnover


class RotorBeta(Rotor):
    """Specialised Rotor; model Beta.

    """

    __encodings = ['L', 'E', 'Y', 'J', 'V', 'C', 'N', 'I', 'X', 'W', 'P', 'B', 'Q',
                   'M', 'D', 'R', 'T', 'A', 'K', 'Z', 'G', 'F', 'U', 'H', 'O', 'S']
    __encodings_rev = ['R', 'L', 'F', 'O', 'B', 'V', 'U', 'X', 'H', 'D', 'S', 'A', 'N',
                       'G', 'Y', 'K', 'M', 'P', 'Z', 'Q', 'W', 'E', 'J', 'I', 'C', 'T']
    __turnover = []

    def __init__(self):
        super().__init__()
        self._name = 'Beta'
        super().__str__()
        self.encodings = RotorBeta.__encodings.copy()
        self.encodings_rev = RotorBeta.__encodings_rev.copy()
        self.turnover = RotorBeta.__turnover


class RotorGamma(Rotor):
    """Specialised Rotor; model Gamma.

    """

    __encodings = ['F', 'S', 'O', 'K', 'A', 'N', 'U', 'E', 'R', 'H', 'M', 'B', 'T',
                   'I', 'Y', 'C', 'W', 'L', 'Q', 'P', 'Z', 'X', 'V', 'G', 'J', 'D']
    __encodings_rev = ['E', 'L', 'P', 'Z', 'H', 'A', 'X', 'J', 'N', 'Y', 'D', 'R', 'K',
                       'F', 'C', 'T', 'S', 'I', 'B', 'M', 'G', 'W', 'Q', 'V', 'O', 'U']
    __turnover = []

    def __init__(self):
        super().__init__()
        self._name = 'Gamma'
        super().__str__()
        self.encodings = RotorGamma.__encodings.copy()
        self.encodings_rev = RotorGamma.__encodings_rev.copy()
        self.turnover = RotorGamma.__turnover


class RotorAbstractFactory:
    """Rotor abstract factory - an implementation of the abstract factory design pattern.

    To enable dynamic creation of whatever flavor of Rotor, so that we can have Enigma machines with however many
    rotors.

    Abstract factory pattern implementation from:

    Chaudhary, M. 2021. Abstract Factory Method – Python Design Patterns [Online]. Uttar Pradesh: GeeksforGeeks.
     Available from: https://www.geeksforgeeks.org/abstract-factory-method-python-design-patterns/
     [Accessed Wed 23 Feb 2022].
    """

    def __init__(self, rotors_factory=None):
        """rotors_factory is the abstract factory

        :param rotors_factory: The Rotor type with which to configure the factory.
        """

        self.rotor_factory = rotors_factory

    def config_factory(self, rotors_factory: Rotor):
        """To enable the factory to be configured to create various specialised Rotor

         - MIRRORS __init__ behaviour (but signature is different) - code smell?

        :param rotors_factory: The Rotor type with which to configure the factory.
        """

        self.rotor_factory = rotors_factory

    def create_rotor(self) -> Rotor:
        """Creates a specialised rotor (of type Rotor) polymorphically using the abstract factory.

        Flavour of Rotor is dependent on the current factory configuration.

        :return: The created Rotor.
        :rtype: Rotor
        """

        created_rotor = self.rotor_factory()
        return created_rotor


class ReflectorSystem:
    """

    """

    # def __init__(self, reflectors_factory=None): # can't get factory pattern working... yet
    def __init__(self, reflector_name='A'):
        """reflector_factory is the abstract factory

        """

        ''' @FIXME can't get this working. for now use the [urgh] pattern adopted in Rotors
        self.reflector_factory = reflectors_factory
        '''
        self.reflector_name = None
        self.reflector_model = None

    def setup(self, reflector_name='A'):
        if type(reflector_name) == str:
            if reflector_name == 'A':
                self.reflector_name = 'A'
                self.reflector_model = ReflectorA()
            elif reflector_name == 'B':
                self.reflector_name = 'B'
                self.reflector_model = ReflectorB()

    def encode(self, letter):
        enc = self.reflector_model.encode(letter)
        print(f'Reflector {self.rotors[-1].__str__()} encoding: {letter}')  # for testing...
        return enc

    ''' part of an experiment to implement factory method design pattern
    def show_reflector(self):
        """ creates and shows reflectors using the abstract factory

        """

        reflector = self.reflector_factory()

        print(f'We have a reflector{reflector}')
        print('it has the following encodings pattern: ')
        print(f'{reflector.encodings}')
    '''


class Reflector(ABC):
    """

    """

    __encodings: list[Any] = list()

    def __init__(self):
        self._name = str
        self._encodings = None

    def __str__(self):
        return self._name

    def encode(self, letter):
        """Encodes a letter using the specific reflector schema in the subclass.

                :param letter: The letter to be encoded.
                :return: The encoded letter.
                """

        if ord(letter) < 65 or ord(letter) > 90:
            raise ValueError
        index_value = ord(letter) - ord('A')

        letter = self._encodings[index_value]
        print(f'Reflector {self.__str__()} encoding: {letter}')  # for testing...

        return letter


class ReflectorA(Reflector):
    """

    """

    __encodings = ['E', 'J', 'M', 'Z', 'A', 'L', 'Y', 'X', 'V', 'B', 'W', 'F', 'C',
                   'R', 'Q', 'U', 'O', 'N', 'T', 'S', 'P', 'I', 'K', 'H', 'G', 'D']

    # have removed param name as workaround to issue described below
    def __init__(self):
        # @TODO trying to be polymorphic, figure out why I cannot call 'self._name = name' from super const.
        # super.__init__(name)

        super().__init__()
        self._name = 'A'
        super().__str__()
        self._encodings = ReflectorA.__encodings.copy()


class ReflectorB(Reflector):
    """

    """

    __encodings = ['Y', 'R', 'U', 'H', 'Q', 'S', 'L', 'D', 'P', 'X', 'N', 'G', 'O',
                   'K', 'M', 'I', 'E', 'B', 'F', 'Z', 'C', 'W', 'V', 'J', 'A', 'T']

    # have removed param name as workaround to issue described below
    def __init__(self):
        # @TODO figure out why I cannot call 'self._name = name' from super const.
        # super.__init__(name)

        super().__init__()
        self._name = 'B'
        super().__str__()
        self._encodings = ReflectorB.__encodings.copy()


class ReflectorC(Reflector):
    """

    """

    __encodings = ['F', 'V', 'P', 'J', 'I', 'A', 'O', 'Y', 'E', 'D', 'R', 'Z', 'X',
                   'W', 'G', 'C', 'T', 'K', 'U', 'Q', 'S', 'B', 'N', 'M', 'H', 'L']

    # have removed param name as workaround to issue described below
    def __init__(self):
        # @TODO figure out why I cannot call 'self._name = name' from super const.
        # super.__init__(name)

        super().__init__()
        self._name = 'C'
        super().__str__()
        self._encodings = ReflectorC.__encodings.copy()


class ReflectorBThin(Reflector):
    """An Enigma reflector.

    Date introduced: 1940
    Model name and number: M4 R1 (M3 + Thin)
    """

    __encodings = ['E', 'N', 'K', 'Q', 'A', 'U', 'Y', 'W', 'J', 'I', 'C', 'O', 'P',
                   'B', 'L', 'M', 'D', 'X', 'Z', 'V', 'F', 'T', 'H', 'R', 'G', 'S']

    # have removed param name as workaround to issue described below
    def __init__(self):
        # @TODO figure out why I cannot call 'self._name = name' from super const.
        # super.__init__(name)

        super().__init__()
        self._name = 'B Thin'
        super().__str__()
        self._encodings = ReflectorBThin.__encodings.copy()


class ReflectorCThin(Reflector):
    """An Enigma reflector.

    Date introduced: 1940
    Model name and number: M4 R1 (M3 + Thin)
    """

    __encodings = ['R', 'D', 'O', 'B', 'J', 'N', 'T', 'K', 'V', 'E', 'H', 'M', 'L',
                   'F', 'C', 'W', 'Z', 'A', 'X', 'G', 'Y', 'I', 'P', 'S', 'U', 'Q']

    # have removed param name as workaround to issue described below
    def __init__(self):
        # @TODO figure out why I cannot call 'self._name = name' from super const.
        # super.__init__(name)

        super().__init__()
        self._name = 'C Thin'
        super().__str__()
        self._encodings = ReflectorCThin.__encodings.copy()


class ReflectorETW(Reflector):
    """An Enigma reflector.

    Date introduced:
    Model name and number: Enigma I
    """

    __encodings = ['A', 'B', 'C', 'D', 'E', 'F', 'G', 'H', 'I', 'J', 'K', 'L', 'M',
                   'N', 'O', 'P', 'Q', 'R', 'S', 'T', 'U', 'V', 'W', 'X', 'Y', 'Z']

    # have removed param name as workaround to issue described below
    def __init__(self):
        # @TODO figure out why I cannot call 'self._name = name' from super const.
        # super.__init__(name)
        super().__init__()
        self._encodings = ReflectorETW.__encodings.copy()


def reflector_factory(name='A'):
    """Factory method to create a reflector of desired type.

    Attempt to implement factory method pattern.

    Intention is to enable creation of whichever reflector, from another object without having to directly create
    instance from class.

    :param name: The name of the intended reflector
    :return: The specified reflector
    """

    reflectors = {'A': ReflectorA,
                  'B': ReflectorB,
                  'C': ReflectorC,
                  'B Thin': ReflectorBThin,
                  'C Thin': ReflectorCThin}

    if name not in reflectors:
        return reflectors['A']()
    else:
        return reflectors[name]()


class ReflectorAbstractFactory:
    """Reflector abstract factory - an implementation of the abstract factory design pattern.

    To enable dynamic creation of whatever flavor of Reflector, so that we can have Enigma machines with however
    many reflectors.

    Abstract factory pattern implementation from:

    Chaudhary, M. 2021. Abstract Factory Method – Python Design Patterns [Online]. Uttar Pradesh: GeeksforGeeks.
     Available from: https://www.geeksforgeeks.org/abstract-factory-method-python-design-patterns/
     [Accessed Wed 23 Feb 2022].
    """

    def __init__(self, reflectors_factory=None):
        """reflectors_factory is the abstract factory

        :param reflectors_factory: The Reflector type with which to configure the factory.
        """

        self.reflector_factory = reflectors_factory

    def config_factory(self, reflectors_factory: Reflector):
        """To enable the factory to be configured to create various specialised Rotor

         - MIRRORS __init__ behaviour (but signature is different) - code smell?

        :param reflectors_factory: The Reflector type with which to configure the factory.
        """

        self.reflector_factory = reflectors_factory

    def create_reflector(self) -> Reflector:
        """Creates a specialised reflector (of type Reflector) polymorphically using the abstract factory.

        Flavour of Reflector is dependent on the current factory configuration.

        :return: The created Reflector.
        :rtype: Reflector
        """

        created_reflector = self.reflector_factory()
        return created_reflector


class Keyboard:
    def __init__(self):
        pass

    def press(self, letter=None):
        if type(letter) != str or (ord(letter.upper()) < 65 or (ord(letter.upper()) > 90)):
            letter = input('Type a letter: ').upper()[0]
            if type(letter) != str or (ord(letter.upper()) < 65 or (ord(letter.upper()) > 90)):
                return ''
        else:
            return letter.upper()


class Enigma:
    def __init__(self):
        self.keyboard = Keyboard()
        self.plugboard = Plugboard()
        self.rotors = Rotors()
        self.reflector = None
        # @TODO should have one abstract factory that has responsibility to create 'families' of objects
        self.rot_af = RotorAbstractFactory()
        self.ref_af = ReflectorAbstractFactory()

    @staticmethod
    def enalpharate_position(position):
        """Helper function to map position value to its equivalent letter in alphabet.

        :param position: The position value to enalpharate.
        :return: The enalpharated position.
        """
        import string
        positions = list(range(1, 27))
        letters = list(string.ascii_uppercase)
        enalpharations = dict(zip(positions, letters))
        return enalpharations[position]

    @staticmethod
    def enumerate_letter(letter):
        """Helper function to map letter to its equivalent position in alphabet.

        :param letter: The position value to enumerate.
        :return: The enumerated position.
        """
        import string
        letter = letter.upper()
        letters = list(string.ascii_uppercase)
        positions = list(range(1, 27))
        enumerations = dict(zip(letters, positions))
        return enumerations[letter]

    def add_reflector(self, reflector_to_add: Reflector):
        """Adds a reflector to this Enigma machine.

        :param reflector_to_add: The reflector to use.
        :type reflector_to_add: Reflector
        """
        self.reflector = reflector_to_add

    def encode(self, letter):
        """The main encode letter use case, synonymous with pressing a key on the Enigma keyboard.

        :param letter: The letter to encode
        :return: The encoded letter.
        """

        plugboard_enc = self.plugboard.encode(letter)
        # @TODO rotations related code smell
        rotors_enc = self.rotors.encode(plugboard_enc)
        reflector_enc = self.reflector.encode(rotors_enc)
        rotors_rev_enc = self.rotors.encode(reflector_enc, True)
        letter_enc = rotors_rev_enc

        return letter_enc


if __name__ == "__main__":
    # case 1
    enigma = Enigma()
    print('Case 1')
    
    # create the desired rotors
    enigma.rot_af.config_factory(RotorI)
    r1 = enigma.rot_af.create_rotor()
    enigma.rot_af.config_factory(RotorII)
    r2 = enigma.rot_af.create_rotor()
    enigma.rot_af.config_factory(RotorIII)
    r3 = enigma.rot_af.create_rotor()

    # configure the ring settings
    r1.set_ring_setting(1)
    r2.set_ring_setting(1)
    r3.set_ring_setting(2)

    # configure the initial rotor positions
    r1.set_position(Enigma.enumerate_letter('A'))
    r2.set_position(Enigma.enumerate_letter('A'))
    r3.set_position(Enigma.enumerate_letter('Z'))

    # add configured rotors to rotors sub-system
    enigma.rotors.add_rotor_to_rotors(r1)
    enigma.rotors.add_rotor_to_rotors(r2)
    enigma.rotors.add_rotor_to_rotors(r3)

    # make and add the reflector
    enigma.ref_af.config_factory(ReflectorB)
    reflector = enigma.ref_af.create_reflector()
    enigma.add_reflector(reflector)

    # perform encryption
    case1 = enigma.encode('A')
    print(f'The result from case 1 is: {case1}')
    print()
    
    # case 2
    enigma = Enigma()
    print('Case 2')
    
    # create the desired rotors
    enigma.rot_af.config_factory(RotorI)
    r1 = enigma.rot_af.create_rotor()
    enigma.rot_af.config_factory(RotorII)
    r2 = enigma.rot_af.create_rotor()
    enigma.rot_af.config_factory(RotorIII)
    r3 = enigma.rot_af.create_rotor()

    # configure the ring settings
    r1.set_ring_setting(1)
    r2.set_ring_setting(1)
    r3.set_ring_setting(2)

    # configure the initial rotor positions
    r1.set_position(Enigma.enumerate_letter('A'))
    r2.set_position(Enigma.enumerate_letter('A'))
    r3.set_position(Enigma.enumerate_letter('A'))

    # add configured rotors to rotors sub-system
    enigma.rotors.add_rotor_to_rotors(r1)
    enigma.rotors.add_rotor_to_rotors(r2)
    enigma.rotors.add_rotor_to_rotors(r3)

    # make and add the reflector
    enigma.ref_af.config_factory(ReflectorB)
    reflector = enigma.ref_af.create_reflector()
    enigma.add_reflector(reflector)

    # perform encryption
    case2 = enigma.encode('A')
    print(f'The result from case 2 is: {case2}')
    print()
    
    # case 3
    enigma = Enigma()
    print('Case 3')
    
    # create the desired rotors
    enigma.rot_af.config_factory(RotorI)
    r1 = enigma.rot_af.create_rotor()
    enigma.rot_af.config_factory(RotorII)
    r2 = enigma.rot_af.create_rotor()
    enigma.rot_af.config_factory(RotorIII)
    r3 = enigma.rot_af.create_rotor()

    # configure the ring settings
    r1.set_ring_setting(1)
    r2.set_ring_setting(1)
    r3.set_ring_setting(2)

    # configure the initial rotor positions
    r1.set_position(Enigma.enumerate_letter('Q'))
    r2.set_position(Enigma.enumerate_letter('E'))
    r3.set_position(Enigma.enumerate_letter('V'))

    # add configured rotors to rotors sub-system
    enigma.rotors.add_rotor_to_rotors(r1)
    enigma.rotors.add_rotor_to_rotors(r2)
    enigma.rotors.add_rotor_to_rotors(r3)

    # make and add the reflector
    enigma.ref_af.config_factory(ReflectorB)
    reflector = enigma.ref_af.create_reflector()
    enigma.add_reflector(reflector)

    # perform encryption
    case3 = enigma.encode('A')
    print(f'The result from case 3 is: {case3}')


```

Output:
```
Case 1

Rotor 3 (III) encoded: B
Rotor 2 (II) encoded: J
Rotor 1 (I) encoded: Z
Reflector B encoding: T
Rotor 1 (I) encoded: L
Rotor 2 (II) encoded: K
Rotor 3 (III) encoded: U
The result from case 1 is: U

Case 2

Rotor 3 (III) encoded: D
Rotor 3 rotated, so input to next is: C
Rotor 2 (II) encoded: D
Rotor 1 (I) encoded: F
Reflector B encoding: S
Rotor 1 (I) encoded: S
Rotor 2 (II) encoded: E
Rotor 3 (III) encoded: C
Rotor 3 rotated, so output is: B
The result from case 2 is: B

Case 3

Rotor 3 (III) encoded: U
Rotor 3 rotated, so input to next is: Y
Rotor 2 (II) encoded: K
Rotor 2 rotated, so input to next is: F
Rotor 1 (I) encoded: B
Rotor 1 rotated, so input to next is: K
Reflector B encoding: N
Rotor 1 (I) encoded: A
Rotor 1 rotated, so input to next is: J
Rotor 2 (II) encoded: Y
Rotor 2 rotated, so input to next is: T
Rotor 3 (III) encoded: H
Rotor 3 rotated, so output is: L
The result from case 3 is: B
```

### Enigma Machine Demonstration
Now in the cell below, demonstrate that you can put all of the elements together. Your full Engima machine should support a plugboard with up to 10 leads, three or four rotors, and a reflector. In addition it should be able to encode an entire string of characters made from the letters `A`-`Z`, correctly advancing the rotors.

Again the code is up to you, and you may want to include additional points to demonstrate your Enigma machine, but please ensure you include the following two examples at minimum:
#### Example 1
Set up your enigma machine with rotors `I II III`, reflector `B`, ring settings `01 01 01`, and initial positions `A A Z`.

The plugboard should map the following pairs: `HL MO AJ CX BZ SR NI YW DG PK`.

*The result of encoding the string* `HELLOWORLD` *should be* `RFKTMBXVVW`.

#### Example 2
Set up your enigma machine with rotors `IV V Beta I`, reflector `A`, ring settings `18 24 03 05`, and initial positions `E Z G P`. 

The plugboard should map the following pairs: `PC XZ FM QA ST NB HY OR EV IU`.

*Find the result of decoding the following string:* `BUPXWJCDPFASXBDHLBBIBSRNWCSZXQOLBNXYAXVHOGCUUIBCVMPUZYUUKHI`.

(You should run this string through your Enigma machine in the cell below; do not just include the result.)

<!-- grade_id: cell-993ccf0331dfd6eb, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part Two – Code Breaking
In this part of the assignment you will be given some ciphertext that has been encrypted using an Enigma machine, an idea of what the original text might contain (a *crib*), and some partial information about the machine. Your goal will be to provide the original plaintext and the full machine settings.

The Bletchley codebreakers were able to combine weaknesses in the machine's encryption, mathematical techniques, and computing power to solve German codes. In this section you only need to use using computing power. You can *brute force* the settings by trying each one until you get the one you are looking for. 

The number of possible settings for Enigma is still too vast to break a code through brute force alone, but you can use the partial information to narrow the search into something feasible on a modern computer. Even on weaker hardware, none of the codes will require more than a few minutes maximum with suitable code.

There is a cell beneath this one for you to include your results for each of the decoded strings. You must include the full decoded message, plus the settings that were missing. You should also point us towards the code (in a separate file) which was used to crack these strings, with instructions for replicating your results. 

Marks are awarded for correct solutions, and for coding style, as outlined in the assignment instructions. Solutions which can be run without manual intervention (e.g. manually changing variables) are likely to receive more credit, but this is not a firm requirement. Any manual steps required must be explained clearly.

### Codes
Each code contains the ciphertext (the encrypted text), and a crib: a word or phrase that you think appears exactly *somewhere* within the original text.

You should use the usual rotors and reflectors specified in the table above unless the question specifies otherwise. So if the question does not specify a rotor, the valid options are only `Beta`, `Gamma`, `I`, `II`, `III`, `IV`, and `V`, the valid options for reflectors are just `A`, `B`, and `C`.

The machines in this section will only ever use 3 rotors – meaning 3 ring settings and 3 starting positions also.

It is possible that more than one set of Enigma machine settings will produce an output containing the crib word. In this case, you must deduce on your own which one is correct based on the output contents.

#### Code 1
You recovered an Enigma machine! Amazingly, it is set up in that day's position, ready for you to replicate in your software. But unfortunately the label has worn off the reflector. All the other settings are still in place, however. You also found a book with the title "SECRETS" which contained the following code, could it be so simple that the code contains that text?

* Code: `DMEXBMKYCVPNQBEDHXVPZGKMTFFBJRPJTLHLCHOTKOYXGGHZ`
* Crib: `SECRETS`


* Rotors: `Beta Gamma V`
* Reflector: Unknown
* Ring settings: `04 02 14`
* Starting positions: `MJM`
* Plugboard pairs: `KI XN FL`

#### Code 2
You leave the machine in the hands of the university. The team have cracked the day's settings thanks to some earlier codebreaking, but unfortunately, the initial rotor positions are changed for each message. For the message below, the team has no idea what the initial settings should be, but know the message was addressed to them. Help them out.

* Code: `CMFSUPKNCBMUYEQVVDYKLRQZTPUFHSWWAKTUGXMPAMYAFITXIJKMH`
* Crib: `UNIVERSITY`


* Rotors: `Beta I III`
* Reflector: `B`
* Ring settings: `23 02 10`
* Starting positions: Unknown
* Plugboard pairs: `VH PT ZG BJ EY FS`

#### Code 3
The department has intercepted a message from the admissions team. They know it contains the word "THOUSANDS" and they are worried it might relate to how many students are arriving next semester. But the admissions team are a bit unusual: they *love* even numbers, and *hate* odd numbers. You happen to know they will never use an odd-numbered rotor, ruling out `I`, `III`, and `V`. They will also never use a *ring setting* that has even a single odd digit: `02` is allowed but `11` is certainly not, and even `12` is banned.

* Code: `ABSKJAKKMRITTNYURBJFWQGRSGNNYJSDRYLAPQWIAGKJYEPCTAGDCTHLCDRZRFZHKNRSDLNPFPEBVESHPY`
* Crib: `THOUSANDS`


* Rotors: Unknown but restricted (see above)
* Reflector: Unknown
* Ring settings: Unknown but restricted (see above)
* Starting positions: `EMY`
* Plugboard pairs: `FH TS BE UQ KD AL`

#### Code 4
On my way home from working late as I walked past the computer science lab I saw one of the tutors playing with the Enigma machine. Mere tutors are not allowed to touch such important equipment! Suspicious, I open the door, but the tutor hears me, and jumps out of the nearest window. They left behind a coded message, but some leads have been pulled out of the machine. It might contain a clue, but I'll have to find the missing lead positions (marked with question marks in the settings below).

* Code: `SDNTVTPHRBNWTLMZTQKZGADDQYPFNHBPNHCQGBGMZPZLUAVGDQVYRBFYYEIXQWVTHXGNW`
* Crib: `TUTOR`


* Rotors: `V III IV`
* Reflector: `A`
* Ring settings: `24 12 10`
* Starting positions: `SWU`
* Plugboard pairs: `WP RJ A? VF I? HN CG BS`

#### Code 5
I later remembered that I had given the tutor permission to use the Enigma machine to solve some codes I'd received via email. As for the window, they are just a big fan of parkour, this is always how they leave the building. It seems they are stuck on one last code. It came in via email so we suspect it's just spam, probably related to a social media website, but you never know when you'll find a gem in that kind of stuff.

The tutor has narrowed the search and found most of the settings, but it seems this code was made with a non-standard reflector. Indeed, there was a photo attached to the email along with the code. It appears that the sender has taken a standard reflector, cracked it open, and swapped some of the wires – two pairs of wires have been modified, by the looks of the dodgy soldering job. 

To be clear, a single wire connects two letters, e.g. mapping `A` to `Y` and `Y` to `A`. The sender has taken two wires (fours pairs of letters), e.g. `A-Y` and `H-J`, and swapped one of the ends, so one option would be `H-Y` and `A-J`. They did this twice, so they modified eight letters total (they did not swap the same wire more than once). 

In your answer, include what the original reflector was and the modifications.

* Code: `HWREISXLGTTBYVXRCWWJAKZDTVZWKBDJPVQYNEQIOTIFX`
* Crib: the name of a social media website/platform


* Rotors: `V II IV`
* Reflector: Unknown and non-standard (see above)
* Ring settings: `06 18 07`
* Starting positions: `AJL`
* Plugboard pairs: `UG IE PO NX WT`

<!-- grade_id: cell-733ec96b80f96616, points: 40, grade: true, solution: true -->
YOUR ANSWER HERE (double click here to edit, write in [Markdown](https://jupyter-notebook.readthedocs.io/en/stable/examples/Notebook/Working%20With%20Markdown%20Cells.html), run to render)

## Advanced Work
Finally, there are a small proportion of marks available to rewards those who push beyond the specification we have presented here, in any way you find interesting.

It is completely acceptable to leave this section blank. The assignment is still plenty of work without doing extra. This assignment is only worth a proportion of the unit, and this section is only worth a proportion of the assignment. We have designed the marking system such that doing well on the other parts of the unit will still be more than enough to get the highest possible classification without submitting anything here.

Academic excellence (the highest possible marks) requires going beyond what you have been directly been taught or asked to do. This will likely become an even bigger factor in later units. This section is an opportunity to demonstrate that ability if you wish.

Please use the text cell below to describe your additional work, pointing to where in your code you demonstrate the work. If you wish to develop your code in a way that would break any of the tests above, you can create a separate folder in your submission for the more advanced version of the code.

Of course, more advanced features will be worth more marks. Your ability to explain your work academically is also important, so consider your presentation style. In particular, considering the programming *theory* of what you are doing (e.g. complexity, mathematical correctness) rather than simply explaining *what* you did is worth more credit. Have fun!

<!-- grade_id: cell-6df98ac7c04bb8f7, points: 20, grade: true, solution: true -->
# OOP design
## Abstract factory design pattern

This project was designed to use the abstract factory design pattern.  The main reason for this was so that Enigma machines with any amount of rotors of any type, could be fitted.  The secondary reason was to capitalise on inheritance and polymorphism, so that each specialised rotor would behave as an actual object of type Rotor.

Also, I wanted to use OOP principals to design the Enigma machine to resemble and work as close to the original as possible.  This turned out to be a lamentable design choice as there simply was not enough time to implement the design as intended.